---
title: Test
description: >-
  Test


---
## Hello, World!

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: a56cc9e540
```

Let's do a Python "Hello World".

`@instructions`
print Hello, World!

`@hint`
use print(string)



`@solution`
```{python}
print("Hello, World!")
```
`@sct`
```{python}
return True
```


